## Convert to Binary

<!-- grade_id: d39095, points: 0, grade: false, solution: true -->
```python
def bin(n):
    result = []
    q = n
    while (q > 0) :
        r = q % 2
        q = q //2
        result.append(r)
    result.reverse()
    return result
    
    raise NotImplementedError()
```

## Testing

<!-- grade_id: 7dd5ce, points: 1, grade: true, solution: false -->
```python
# Check that bin function returns the correct output for several inputs:
from nose.tools import assert_equal
assert_equal(bin(34), [1, 0, 0, 0, 1, 0])
assert_equal(bin(42), [1, 0, 1, 0, 1, 0])
assert_equal(bin(121), [1, 1, 1, 1, 0, 0, 1])



```

```python

```