# GCD program

<!-- grade_id: 1eb55a, points: 0, grade: false, solution: true -->
```python
def mygcd(b,r):  # modify function name and parameters
    """
    Compute the GCD of positive integers b and r
    """
    # YOUR CODE HERE
    if r == 0:
        return b
    else:
        r2 = b % r
        return(mygcd(r,r2))

    raise NotImplementedError()
```

Now, test your code on a few examples.

```python
mygcd(27,11)
```

Output:
```
1
```

```python
mygcd(81,15)
mygcd(27,9)
```

Output:
```
9
```

<!-- grade_id: a15fc2, points: 3, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
from nose.tools import assert_equal
assert_equal(mygcd(13013,121 ), 11)
assert_equal(mygcd(220,66 ), 22)
assert_equal(mygcd(13,17 ), 1)


```

```python
assert_equal(mygcd(27, 11), 1)
assert_equal(mygcd(81, 15), 3)
assert_equal(mygcd(27, 9), 9)
```

```python

```